# <span style="color:red">Before you turn this in, make sure everything runs as expected.</span>

1. **RESTART THE KERNEL** – in the menubar, select Kernel$\rightarrow$Restart
2. **RUN ALL CELLS** – in the menubar, select Cell$\rightarrow$Run All
3. **VALIDATE THE NOTEBOOK** – in the menubar, click the Validate button

## <span style="color:blue">How to Answer Questions</span>

### <span style="color:blue">Python code answers</span>

Enter your answer any place that says
```python
# Enter your code here
```
<span style="color:red">**AND delete the text.**</span>
```python
raise NotImplementedError # No Answer - remove if you provide an answer
```

### <span style="color:blue">Written answers</span>

Enter your answer any place that says
```
YOUR ANSWER HERE.
```

```python
ANUID = "u7522927"
```

---

# Python quiz 3

This quiz is worth 1% of the course.

## Q 1 – simple function

Write a function called `squared` that takes one numerical argument and returns the square of it (e.g. $f(i)=i^2$)

<!-- grade_id: squared-defn, points: 0, grade: false, solution: true -->
```python
# complete this function
def squared(x):
    y=x*x
    return y

```

<!-- grade_id: squared-defined, points: 0.05, grade: true, solution: false -->
```python
# This part worth 0.05

# This cell is testing your code
"""Q1 simple function : correct name and callable"""
assert "squared" in locals()
assert callable(squared), "it's not a function!"
```

<!-- grade_id: squared-val, points: 0.05, grade: true, solution: false -->
```python
# This part worth 0.05

# This cell is testing your code
"""Q1 simple function : correct value"""
```

Output:
```
'Q1 simple function : correct value'
```

---

## Q2 – converting a data line to a dict

In the Donor by Cancer type file, the first line of the file is the column names and the remainder are data lines.

The objective is to write a function that takes one data line and returns a `dict` keyed by the column names with values  from the corresponding data line fields. To do this you need to provide the function with the field names and the data line to be converted. The line to be converted needs to be cleaned up and split on the delimiting character.

For example
```python
line_to_fields(("Individual Id", "Project ID"), "D123\tNasty-Cancer\n")
```
should return
```python
{'Individual Id': 'D123', 'Project ID': 'Nasty-Cancer'}
```

I will test using synthetic examples and the one given below.

```python
field_names = ("Donor_ID", "Project_Code")
line = "DO36068\tCNS-PiloAstro\n"
```

<!-- grade_id: delimiter-soln, points: 0, grade: false, solution: true -->
```python
# complete this function
import re
def line_to_fields(fields: list[str], data_line: str) -> dict[str, str]:
    result = {}
    linearr=re.split('\t|\n',data_line)
    del linearr[-1]
    print(linearr)
    i = 0
    for elem in fields:
        result[elem]=linearr[i]
        i=i+1
    return result
print(line_to_fields(("Individual Id", "Project ID"), "D123\tNasty-Cancer\n"))


```

Output:
```
['D123', 'Nasty-Cancer']
{'Individual Id': 'D123', 'Project ID': 'Nasty-Cancer'}
```

<!-- grade_id: func_signature, points: 0.05, grade: true, solution: false -->
```python
# This part worth 0.05

# This cell is testing your code
"""Q2 line_to_fields() : correct variable names and modules"""
from gutils import check

check.allowed_modules()
check.expected_variables_exist(["line_to_fields", "line", "field_names"], locals())
```

Output:
```
Cell In[16], line 8
      5 from gutils import check
      7 check.allowed_modules()
----> 8 check.expected_variables_exist(["line_to_fields", "line", "field_names"], locals())

File /opt/conda/envs/python3.10/lib/python3.10/site-packages/gutils/check.py:28, in expected_variables_exist(var_names, scope, callables)
     25     msgs.append(f"The following are not callable: {msg}")
     27 if msgs:
---> 28     raise AssertionError("\n".join(msgs))

AssertionError: The following expected variables are missing: line, field_names
```

<!-- grade_id: delimiter-provided, points: 0.05, grade: true, solution: false -->
```python
# This part worth 0.05

# This cell is testing your code
"""Q2 line_to_fields() : correct values on provided data line"""
```

Output:
```
'Q2 line_to_fields() : correct values on provided data line'
```

<!-- grade_id: delimiter-other, points: 0.05, grade: true, solution: false -->
```python
# This part worth 0.05

# This cell is testing your code
"""Q2 line_to_fields() : correct on another value"""
```

Output:
```
'Q2 line_to_fields() : correct on another value'
```

## Q3 – useful data from a file

Write a parser for the file at `path` (it is the one used in Exercise 5.30(4)), generating a list of dicts using your `line_to_fields()` function.

```python
path = "data/donor_by_cancer_type.tsv"
```

<!-- grade_id: parse-disease-soln, points: 0, grade: false, solution: true -->
```python
# complete this function
def parse_project_data(path) -> list[dict]:
    result = []
    file1 = open(path, "r")
    fieldnames = file1.readline()
    header=re.split('\t|\n',fieldnames)
    del header[-1]
    lines = file1.readlines()
    for data_line in lines:
        result.append(line_to_fields(header, data_line))
    return result


```

<!-- grade_id: parse-disease-checks, points: 0.05, grade: true, solution: false -->
```python
# This part worth 0.05

# This cell is testing your code
"""Q3 parse_project_data() : basic checks"""
check.allowed_modules()
check.expected_variables_exist(["parse_project_data"], locals())
assert callable(parse_project_data), f"{parse_project_data} is not a function"
```

<!-- grade_id: parse-disease-num_lines, points: 0.1, grade: true, solution: false -->
```python
# This part worth 0.1

"""Q3 parse_project_data() : correct number of lines"""
# Additional checks.
```

Output:
```
'Q3 parse_project_data() : correct number of lines'
```

<!-- grade_id: parse-disease-, points: 0.1, grade: true, solution: false -->
```python
# This part worth 0.1

# This cell is testing your code
"""Q3 parse_project_data() : exactly correct values"""
```

Output:
```
'Q3 parse_project_data() : exactly correct values'
```

## Q 4 – counting types

Having generated a useful data structure from the data file, we now want to compute some summary statistics. Specifically, we want to compute how many donors are registered for each cancer.

For example
```python
data = [
    {'Donor_ID': 'DO220887', 'Project_Code': 'Skin-Melanoma'},
    {'Donor_ID': 'DO220872', 'Project_Code': 'Skin-Melanoma'},
    ]
count_cancer_types(data)
```
should return
```python
{'Skin-Melanoma': 2}
```

<!-- grade_id: count-types-soln, points: 0, grade: false, solution: true -->
```python

from collections import Counter
def count_cancer_types(records: list[dict]) -> dict[str, int]:
    result = {}
    for record in records:
        code = record["Project_Code"]
        if code in result:
            result[code] += 1
        else:
            result[code] = 1
        # or replace above conditional with this line
        # result[code] = result.get(code, 0) + 1
    return result

```

<!-- grade_id: count-types-check, points: 0.05, grade: true, solution: false -->
```python
# This part worth 0.05

# This cell is testing your code
"""Q3 count_cancer_types() : basic checks"""
check.allowed_modules()
check.expected_variables_exist(["count_cancer_types"], locals())
assert callable(parse_project_data), f"{parse_project_data} is not a function"
```

<!-- grade_id: count-types-sample, points: 0.15, grade: true, solution: false -->
```python
# This part worth 0.15

# This cell is testing your code
"""Q4 count_cancer_types() : correct values for sample data"""
check.expected_variables_exist(["count_cancer_types"], locals())
assert callable(count_cancer_types), f"{get_numeric} is not a function"
# Additional checks.
```

<!-- grade_id: count-types-full, points: 0.3, grade: true, solution: false -->
```python
# This part worth 0.3

# This cell is testing your code
"""Q4 count_cancer_types() : correct values on other cases"""
```

Output:
```
'Q4 count_cancer_types() : correct values on other cases'
```

```python

```